Before you turn this homework in, make sure everything runs as expected. First, **restart the kernel** and then **run all cells** (in the menubar, select Kernel$\rightarrow$Restart \& Run all). If you get error messages, the grader will get the same errors and won't be able to grade your work. 

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE". You may delete the line(s) that say "**raise** NotImplementedError()" and replace them with your actual work. Write all your work in the cells provided in this notebook. **DO NOT** add or delete any cells. 

Please fill in your name in the cell below.

```python
NAME = "Kevin Pham"
```

---

# Homework 3

In all problems below, please comment your code sufficiently well so that the grader can follow what you are doing with ease. For non-coding answers, please make sure to formulate your explanation and answers in the form of complete English sentences. 

## Problem 1 (a)

Use the code given below to create a sequence of n Bernoulli trials with success probabilty p per trial. In a Bernoulli trial, a 1 counts as a success. A string of consecutive successes is known as a success run. Write a function (named "count_runs") that takes the sequence of Bernoulli trials as its input and returns the counts for runs of length k for each k observed in a dictionary.

Example: n=13 Bernoulli trials [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 1, 0, 0] were simulated with p=0.5. There are 3 runs of length k=1, and 1 run of length k=2. Hence, your code should return {1:3, 2:1}. Your code should work even with a different seed or with different values of n and/or p. 

```python
import random

random.seed(10)
n=13
p=0.5
seq = random.choices([1, 0],[p,1-p],k=n) # create sequence of n Bernoulli trials
print(seq)
```

Output:
```
[0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 1, 0, 0]
```

<!-- grade_id: cell-631ba149fea50a1d, points: 0, grade: false, solution: true -->
```python
# COMPLETE
def run_counts(seq):  # define function.
   
    # Between each '0's are going to be combined strings of 1's.
    final_dict = {} # final dict. for output
    trim_list = [] # trimmed list that will exclude the 0's
    # 'Joins' consecutive 1's into a single string per entry of a list while excluding 0's
    str_list = ''.join([str(value) for value in seq]).split('0') 
    
    [trim_list.append(value) for value in str_list if value != ''] # trim down the list to remove the empty spaces
    adj_set = list(set(trim_list))[::-1] # To traverse only the unique keys of the dictionary
    
    for i in adj_set: # Start queuing in the dictionary
        final_dict[len(i)] = trim_list.count(i) # queue in by length of the string per unique key
    
    return(final_dict) # return the dictionary
    
print(run_counts(seq[1]))

    
    # YOUR CODE HERE
    # raise NotImplementedError()
```

Output:
```
Input In [7], in <cell line: 18>()
     14         final_dict[len(i)] = trim_list.count(i) # queue in by length of the string per unique key
     16     return(final_dict) # return the dictionary
---> 18 print(run_counts(seq[1]))

Input In [7], in run_counts(seq)
      6 trim_list = [] # trimmed list that will exclude the 0's
      7 # 'Joins' consecutive 1's into a single string per entry of a list while excluding 0's
----> 8 str_list = ''.join([str(value) for value in seq]).split('0') 
     10 [trim_list.append(value) for value in str_list if value != ''] # trim down the list to remove the empty spaces
     11 adj_set = list(set(trim_list))[::-1] # To traverse only the unique keys of the dictionary

TypeError: 'int' object is not iterable
```

<!-- grade_id: cell-86f7e5fef8b7d4dd, points: 3, grade: true, solution: false -->
```python
# Check if your function works by executing the code below. 
# If the assert statements below DO NOT produce any errors, your code works on the test cases. 

assert run_counts([0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 1, 0, 0])=={1:3,2:1}
assert run_counts([1, 0, 0, 1, 1, 0, 1, 0, 0, 1, 0, 0, 0, 0, 1, 1, 1, 0, 1, 1, 0, 1]) == {1: 4, 2: 2, 3: 1}

```

## Problem 1 (b)

Continuing with the Bernoulli trials from the previous problem. We want to numerically estimate the probability to have at least one run of length 5 or longer for n=100 and p=0.5. Write code to estimate this probability with two after decimal digits of accuracy and report your result in form of a print statement of the following form: 
"The estimated probability to see at least one run of length (run_length) in (n) trials with success probability (p) is (estimated_probability)". Where you code should automatically supply the values of run_length, n, p, and the estimated_probability. 

<!-- grade_id: cell-24926ec4f814a84e, points: 2, grade: true, solution: true -->
```python
# COMPLETE
n = 100 # trials
p = 0.5 # probability per trial
num = 10000 # total number of iterations

# empty list for the sequence of bernoulli trials
trials = []
# empty list to append the successes
result = []

for i in range(num): # start traversing through each iterations
    trials.append(random.choices([1, 0],[p,1-p],k=n))
    if(max(run_counts(trials[i]))>=5): # check each trial to check if it's length >= 5
        result.append(True)

print("The estimation for n = 100 with p = 0.5 with length 5", sum(result)/num)



# raise NotImplementedError()
```

Output:
```
The estimation for n = 100 with p = 0.5 with length 5 0.8129
```

# Problem 1 (c) 

What is more likely - to see at least one run of length 5 or longer in 100 trials with p=0.5 or to see at least one run of length 7 or longer in 100 trials with p=0.7? Explain how you come to your conclusion. Justify any choices you make in your simulation code. 

<!-- grade_id: cell-99194fcbaaef3c27, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
# checking with p = 0.7 to compare with p = 0.5
p = 0.7 # set new probability

# reset the list of bernoulli trials
trials = []
# reset the list for appended successes
result = []

for i in range(num):
    trials.append(random.choices([1, 0],[p,1-p],k=n))
    if(max(run_counts(trials[i]))>=5): # check each trial to check if it's length >= 5
        result.append(True)

print("Estimation with n = 100 and p = 0.7 of length 7 is",sum(result)/num)

# using the same n = 100 and iterations as before
# raise NotImplementedError()
```

Output:
```
Estimation with n = 100 and p = 0.7 of length 7 is 0.9992
```

<!-- grade_id: cell-2d53817c1f1588a8, points: 1, grade: true, solution: true -->
Computing the above results, it can be seen that it is more likely to get a length of 7 with p = 0.7 vs. its predecessor being length 5 for p = 0.5. Even using an iteration as high as 10000, the probability remains stable during repeats of the program.

## Problem 2

A stick of length 1 is broken randomly along its length into two pieces. This means that the break point is a Uniform(0,1) random variable. What is the probability that the longer piece is more than twice as long as the shorter piece? 

a) Write a simulation that approximates the answer to this problem to two after decimal digit accuracy. 

**Note:** you should *not* use your knowledge of the theoretical answer in your simulation. That defies the purpose of a simulation. Simulate a uniform random variable to represent the break-point, instead. And actually compare the length of the "pieces" in your code. 

<!-- grade_id: cell-eedec05f0fd650d1, points: 3, grade: true, solution: true -->
```python
# COMPLETE
stick_list = []
n = 30000

for i in range(n):
    break_pt = random.uniform(0, 1) # break the stick of length 1 in a random spot
    if((break_pt/(1-break_pt)>2 or (1-break_pt)/break_pt>2)): # check to see which of the two sticks is longer
        stick_list.append(True) # append the success


print("For n =",n,"the approximation is",round(sum(stick_list)/n,2)) # round to 2 decimal digits
# raise NotImplementedError()
```

Output:
```
For n = 30000 the approximation is 0.67
```

b) Find the actual theoretical answer (using your Math 161A probability knowledge). Provide a short explanation for how you derive your answer. 

<!-- grade_id: cell-b03e8954539eddc8, points: 1, grade: true, solution: true -->
From 161A, If one interval is more than twice as long as the other, then it means that one of the two intervals is longer than 2/3 (and the other is less than 1/3).Letting 'A' be the event of one interval being greater than 2/3 then the complement of the event 'A' (A') is that neither interval is greater than 2/3, which also means neither interval is less than 1/3 as well. Since this falls under a uniform distribution, the probability of A' being the probability that the break point falls between 1/3 and 2/3, would be exactly 1/3. Therefore P(A) = 1 - P(A') = 2/3.

## Problem 3

Use the random package to simulate rolling five fair, six-sided dice. We'll call the result a "full house" if two and three of the dice, respectively show the same face. Example: 2,4,2,4,4 is a full house. 5,5,5,5,5 is not a full house. 

(a) Using your Math 161A knowledge, find the theoretical probability of obtaining a "full house", when rolling five fair, six-sided dice. Show your work in deriving your answer. 

Hint: Latex notation for combinations is "\binom{n}{k}" to produce $\binom{n}{k}$. Wrap Latex math code in dollar symbols (\$) to compile it. 

<!-- grade_id: cell-c185f3c9fb9b3cb9, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

(b) Write a function (named "count_rolls()") that takes the result of simulating rolling five fair, six-sided dice and returns a dictionary that keeps track of how often each number appeared.

Example: If the rolls were 4,3,4,2,5 then your dictionary could be {4: 2, 3: 1, 2: 1, 5: 1}. I do not care about the order in which the keys appear in the dictionary. An equally good answer would be {2:1, 3:1, 4:2, 5:1}, for example.

<!-- grade_id: cell-358bc28873f8d117, points: 0, grade: false, solution: true -->
```python
# COMPLETE
import random 
random.seed(10)                               # feel free to change or remove random seed
                                              # this random seed corresponds to the example described above. 

result = random.choices(range(1,7),[1]*6,k=5) # simulate five die rolls
print(result)                                 # look at the rolls (feel free to delete later)

def count_rolls(result):
    die_dict = {} # empty dictionary for result
    for i in range(len(result)): # traverse and add each key corresponding to what side each die lands on
        if result[i] in die_dict: # add each dict key by side of the die facing up
            die_dict[result[i]] += 1 # increment for each respective side facing up
        else:
            die_dict[result[i]] = 1 # adding a new side to the dictionary 
    return(die_dict) # return the dict when done

dexample = count_rolls(result)
count_rolls(result)
# YOUR CODE HERE
# raise NotImplementedError()
```

Output:
```
[4, 3, 4, 2, 5]
{4: 2, 3: 1, 2: 1, 5: 1}
```

<!-- grade_id: cell-df2878ab73ef6163, points: 2, grade: true, solution: false -->
```python
# Check if your function works by executing the code below. 
# If the assert statements below DO NOT produce any errors, your code works on the test cases. 

assert count_rolls([4, 3, 4, 2, 5]) == {4: 2, 3: 1, 2: 1, 5: 1}
assert count_rolls([3, 4, 6, 3, 4]) == {3: 2, 4: 2, 6: 1}

```

(c) Simulate rolling five fair, six-sided dice n=10,000 times. Use your simulation results to estimate the probability of a full house with three after decimal digits accuracy. 

<!-- grade_id: cell-deb5efb0ac393d4d, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
n = 10000
full_house = 0
for i in range(n): # traverse through each trial
    roll = count_rolls(random.choices(range(1,7),[1]*6,k=5)) # run the above die simulator
    if (len(roll) == 2 & 2 in roll.values()):
        full_house += 1 # increment each time a full house has been satisfied

print("The estimated probability based on n =",n,"iterations is",round(full_house/n,3)) # final statement with 3 decimal digit accuracy
#raise NotImplementedError()
```

Output:
```
The estimated probability based on n = 10000 iterations is 0.038
```

## Problem 4

Consider the (uncommented) mystery function defined below whose input d is a dictionary. Assume that all of the keys of the input dictionary are strings and all values of the input dictionary are integers. 

```python
def mystery(d):
    r = len(d)
    val = list(d.values())
    keys = list(d.keys())

    zeroes= []
    for i in range(r):
        if not val[i]:
            zeroes.append(i)

    for zero in zeroes:
        del d[keys[zero]]
    
    r = len(d)
    val = list(d.values())
    keys = list(d.keys())
    dd = {}

    for i in range(r):
        if val[i]//2 == val[i]/2:
            keys[i] = keys[i].upper()
        else:
            keys[i] = keys[i].lower()

        dd[keys[i]] = val[i]
    
    return dd
```

(a) Summarize what the mystery function does in a few short succinct English sentences. A line-by-line translation of Python code into English is *not* a succinct summary. 

<!-- grade_id: cell-1c23f3637a563602, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

(b) Define an input dictionary d below of length at least 5, so that the input d is not identical to mystery(d) and so that mystery(d) is exactly `{'AAA': 4, 'bb':9, 'CCC':20}`

<!-- grade_id: cell-a539720e981dd776, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(c) Instead of the complicated code used to define the mystery function above, write a one-line dictionary comprehension statement that achieves the same effect when applied to a dictionary `d` that has string keys and integer values.  

<!-- grade_id: cell-df433bce48997a2d, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```