# Solution for Homework 3

In all problems below, please comment your code sufficiently well so that the grader can follow what you are doing with ease. For non-coding answers, please make sure to formulate your explanation and answers in the form of complete English sentences. 

## Problem 1 (a)

Use the code given below to create a sequence of n Bernoulli trials with success probabilty p per trial. In a Bernoulli trial, a 1 counts as a success. A string of consecutive successes is known as a success run. Write a function (named "count_runs") that takes the sequence of Bernoulli trials as its input and returns the counts for runs of length k for each k observed in a dictionary.

Example: n=13 Bernoulli trials [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 1, 0, 0] were simulated with p=0.5. There are 3 runs of length k=1, and 1 run of length k=2. Hence, your code should return {1:3, 2:1}. Your code should work even with a different seed or with different values of n and/or p. 

```python
import random

random.seed(10)
n=13
p=0.5
seq = random.choices([1, 0],[p,1-p],k=n) # create sequence of n Bernoulli trials
print(seq)
```

Output:
```
[0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 1, 0, 0]
```

<!-- grade_id: cell-631ba149fea50a1d, points: 0, grade: false, solution: true -->
```python
def run_counts(seq):  
    
    counts = {}                                 # make empty dictionary
        # the run lengths will be the keys and their counts will be the values
    bernoulli_string = "".join([str(i) for i in seq]) 
        # make outcomes into string of 0 and 1
    run_list = bernoulli_string.split('0')            
        # split by 0's - to separate runs of 1's
    run_list = [item for item in run_list if len(item)>0] 
        # get rid of "empty runs"         

    for run in run_list:                         # iterate over all runs
        counts[len(run)] = counts.get(len(run),0) + 1   
            # increase count of each run of the same length you come across by 1
    return counts
```

<!-- grade_id: cell-86f7e5fef8b7d4dd, points: 3, grade: true, solution: false -->
```python
# Check if your function works by executing the code below. 
# If the assert statements below DO NOT produce any errors, your code works on the test cases. 

assert run_counts([0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 1, 0, 0])=={1:3,2:1}
assert run_counts([1, 0, 0, 1, 1, 0, 1, 0, 0, 1, 0, 0, 0, 0, 1, 1, 1, 0, 1, 1, 0, 1]) == {1: 4, 2: 2, 3: 1}

### BEGIN HIDDEN TESTS
"""Check that function works as it is supposed to"""
from nose.tools import assert_equal
assert_equal(run_counts([1, 0, 0, 1, 1, 0, 1, 0, 0, 1, 0, 0, 0, 0, 1]),{1: 4, 2: 1})
assert_equal(run_counts([1, 1, 0, 1, 1, 0, 1, 0, 1, 1, 0, 0, 1, 1, 1, 1, 1, 0, 0, 0, 0, 1, 0, 1, 0]),{2: 3, 1: 3, 5: 1})
### END HIDDEN TESTS
```

## Problem 1 (b)

Continuing with the Bernoulli trials from the previous problem. We want to numerically estimate the probability to have at least one run of length 5 or longer for n=100 and p=0.5. Write code to estimate this probability with two after decimal digits of accuracy and report your result in form of a print statement of the following form: 
"The estimated probability to see at least one run of length (run_length) in (n) trials with success probability (p) is (estimated_probability)". Where you code should automatically supply the values of run_length, n, p, and the estimated_probability. 

<!-- grade_id: cell-24926ec4f814a84e, points: 2, grade: true, solution: true -->
```python
random.seed(12)
n=100           # length of sequence to be generated
p=0.5           # success probability
r=5             # run length
m = 10000       # number of sequences to be generated
long_run = []

for i in range(m):
    seq = random.choices([1, 0],[p,1-p],k=n)  # generate one Bernoulli sequence of length n
    run_dict = run_counts(seq)                # count run lengths
    run_dict = {key:value for (key,value) in run_dict.items() if key>=r} # identify all runs of length >r
    long_run.append(sum(run_dict.values())>0)  
        # checks if there is at least one run of length 5 or longer
        # and appends the Boolean to the long_run list
        
print("The estimated probability to see at least one run of length",r,"or longer \
in",n,"trials with success probability",p, "is:", round(sum(long_run)/m,2))
```

Output:
```
The estimated probability to see at least one run of length 5 or longer in 100 trials with success probability 0.5 is: 0.81
```

# Problem 1 (c) 

What is more likely - to see at least one run of length 5 or longer in 100 trials with p=0.5 or to see at least one run of length 7 or longer in 100 trials with p=0.7? Explain how you come to your conclusion. Justify any choices you make in your simulation code. 

<!-- grade_id: cell-99194fcbaaef3c27, points: 0, grade: true, solution: true -->
```python
random.seed(12)
n=100
p=0.7
r=7
m = 10000
long_run = []

for i in range(m):
    seq = random.choices([1, 0],[p,1-p],k=n) # generate one Bernoulli sequence of length n
    run_dict = run_counts(seq)               # count run lengths
    run_dict = {key:value for (key,value) in run_dict.items() if key>=r} # identify all runs of length >r
    long_run.append(sum(run_dict.values())>0) 
        # checks if there is at least one run of length 5 or longer
        # and appends the Boolean to the long_run list
        
print("The estimated probability to see at least one run of length",r,"or longer \
in",n,"trials with success probability",p, "is:", round(sum(long_run)/m,2))
### END SOLUTION
```

Output:
```
The estimated probability to see at least one run of length 7 or longer in 100 trials with success probability 0.7 is: 0.95
```

<!-- grade_id: cell-2d53817c1f1588a8, points: 1, grade: true, solution: true -->
Run the code a couple of times with reasonable large m (I used m=10,000) to see how much the results change across runs. If they change a lot, increase the value of m. Once you're satisfied that results are fairly stable run the code with (p=0.5 and r=5) as well as with (p=0.7 and r=7). The probability to see at least one run of length 5 in 100 Bernoulli's with p=0.5 is about 0.81. The probability to see at least one run of length 7 in 100 Bernoulli's with p=0.7 is about 0.95. Hence, the p=0.7 run of length 7 is more likely. 

## Problem 2

A stick of length 1 is broken randomly along its length into two pieces. This means that the break point is a Uniform(0,1) random variable. What is the probability that the longer piece is more than twice as long as the shorter piece? 

a) Write a simulation that approximates the answer to this problem to two after decimal digit accuracy. 

**Note:** you should *not* use your knowledge of the theoretical answer in your simulation. That defies the purpose of a simulation. Simulate a uniform random variable to represent the break-point, instead. And actually compare the length of the "pieces" in your code. 

<!-- grade_id: cell-eedec05f0fd650d1, points: 3, grade: true, solution: true -->
```python
### BEGIN SOLUTION 
import random as random 
#random.seed(10)
n=100000
prob = []

for i in range(n):
    bp = random.uniform(0,1) # generate break-point
    pieces = [bp, 1-bp]
    prob.append(max(pieces)>2*min(pieces))
    
print("The simulated value of the probability \
that the longer piece is at least twice as long\
as the shorter piece is", round(sum(prob)/n,2))
```

Output:
```
The simulated value of the probability that the longer piece is at least twice as longas the shorter piece is 0.67
```

b) Find the actual theoretical answer (using your Math 161A probability knowledge). Provide a short explanation for how you derive your answer. 

<!-- grade_id: cell-b03e8954539eddc8, points: 1, grade: true, solution: true -->
To have the longer piece be at least twice as long as the shorter piece, the break-point has to lie either in the interval $[0,1/3)$ - in which case the long piece will be on the right, or the break point has to be in the interval $(2/3,1]$ in which case the long piece is on the left. The probability that a uniform(0,1) random variable falls into an interval $(a,b) \subset [0,1]$ is $b-a$. Hence, the probability that the break-point falls into either interval above is $\frac{1}{3} + \frac{1}{3} = \frac{2}{3}$.

## Problem 3

Use the random package to simulate rolling five fair, six-sided dice. We'll call the result a "full house" if two and three of the dice, respectively show the same face. Example: 2,4,2,4,4 is a full house. 5,5,5,5,5 is not a full house. 

(a) Using your Math 161A knowledge, find the theoretical probability of obtaining a "full house", when rolling five fair, six-sided dice. Show your work in deriving your answer. 

Hint: Latex notation for combinations is "\binom{n}{k}" to produce $\binom{n}{k}$. Wrap Latex math code in dollar symbols (\$) to compile it. 

<!-- grade_id: cell-c185f3c9fb9b3cb9, points: 1, grade: true, solution: true -->
Since all dice are fair, all possible outcomes are equally likely. In this case, to compute a probability, you could count favorables (number of different ways in which you can make a "full house") and possibles (number of different ways in which you can toss five dice). There are $6^5$ different ways to toss five six-sided dice (we're making order of the dice matter, here) and there are $6\binom{5}{3}5\binom{2}{2}$ ways to create "full houses". That's because you have six choices for the number that appears three times (and then $\binom{5}{3}$ choices of where to put these numbers - remember that order mattters) and then you have 5 choices (you can't pick the number that was already picked) to fill in the remaining two empty slots. Hence 

$$ \mbox{Full House} = \frac{6\binom{5}{3}5}{6^5} = \frac{300}{7776} = 0.0386$$

(b) Write a function (named "count_rolls()") that takes the result of simulating rolling five fair, six-sided dice and returns a dictionary that keeps track of how often each number appeared.

Example: If the rolls were 4,3,4,2,5 then your dictionary could be {4: 2, 3: 1, 2: 1, 5: 1}. I do not care about the order in which the keys appear in the dictionary. An equally good answer would be {2:1, 3:1, 4:2, 5:1}, for example.

<!-- grade_id: cell-358bc28873f8d117, points: 0, grade: false, solution: true -->
```python
import random 
random.seed(10)                               # feel free to change or remove random seed
                                              # this random seed corresponds to the example described above. 

result = random.choices(range(1,7),[1]*6,k=5) # simulate five die rolls
print(result)                                 # look at the rolls (feel free to delete later)

def count_rolls(result):
    
    counts = {}                                 # make empty dictionary
    # the rolls will be the keys and their counts will be the values
    for roll in result:                         # iterate over all five rolls
        counts[roll] = counts.get(roll,0) + 1   # increase count of each roll you come across by 1
    
    return counts
```

Output:
```
[4, 3, 4, 2, 5]
```

<!-- grade_id: cell-df2878ab73ef6163, points: 2, grade: true, solution: false -->
```python
# Check if your function works by executing the code below. 
# If the assert statements below DO NOT produce any errors, your code works on the test cases. 

assert count_rolls([4, 3, 4, 2, 5]) == {4: 2, 3: 1, 2: 1, 5: 1}
assert count_rolls([3, 4, 6, 3, 4]) == {3: 2, 4: 2, 6: 1}

### BEGIN HIDDEN TESTS
"""Check that function works as it is supposed to"""
from nose.tools import assert_equal
assert count_rolls([4, 5, 6, 2, 1]) == {4: 1, 5: 1, 6: 1, 2: 1, 1: 1}
assert count_rolls([4, 2, 1, 1, 4]) == {4: 2, 2: 1, 1: 2}
assert count_rolls([5, 3, 2, 2, 5]) == {5: 2, 3: 1, 2: 2}
### END HIDDEN TESTS
```

(c) Simulate rolling five fair, six-sided dice n=10,000 times. Use your simulation results to estimate the probability of a full house with three after decimal digits accuracy. 

<!-- grade_id: cell-deb5efb0ac393d4d, points: 2, grade: true, solution: true -->
```python
import random 
random.seed(10)
n=10000
full_house = []

for i in range(n):
    result = random.choices(range(1,7),[1]*6,k=5)    # simulate five rolls
    counts = count_rolls(result)                     # count how often each number appears
    full_house.append(set(counts.values()) == {2,3}) # check whether result is full house and save Boolean
    
print("The probability of a full house is estimated to be ", round(sum(full_house)/n,3))
```

Output:
```
The probability of a full house is estimated to be  0.038
```

## Problem 4

Consider the (uncommented) mystery function defined below whose input d is a dictionary. Assume that all of the keys of the input dictionary are strings and all values of the input dictionary are integers. 

```python
def mystery(d):
    r = len(d)
    val = list(d.values())
    keys = list(d.keys())

    zeroes= []
    for i in range(r):
        if not val[i]:
            zeroes.append(i)

    for zero in zeroes:
        del d[keys[zero]]
    
    r = len(d)
    val = list(d.values())
    keys = list(d.keys())
    dd = {}

    for i in range(r):
        if val[i]//2 == val[i]/2:
            keys[i] = keys[i].upper()
        else:
            keys[i] = keys[i].lower()

        dd[keys[i]] = val[i]
    
    return dd
```

(a) Summarize what the mystery function does in a few short succinct English sentences. A line-by-line translation of Python code into English is *not* a succinct summary. 

<!-- grade_id: cell-1c23f3637a563602, points: 2, grade: true, solution: true -->
Key-value pairs for which the value is zero are removed from the dictionary. Keys with even values are made all upper case. Keys with odd values are made all lower case. 

(b) Define an input dictionary d below of length at least 5, so that the input d is not identical to mystery(d) and so that mystery(d) is exactly `{'AAA': 4, 'bb':9, 'CCC':20}`

<!-- grade_id: cell-a539720e981dd776, points: 1, grade: true, solution: true -->
```python
d = {'aAa':4, 'Bb': 9, 'ccC': 20, 'D':0, 'e': 0}

# Fill in the dictionary d above with keys and values so that the three statements below all evaluate to True. 
print(len(d) >= 5)
print(mystery(d) != d)
print(mystery(d) == {'AAA': 4, 'bb':9, 'CCC':20})
```

Output:
```
True
True
True
```

(c) Instead of the complicated code used to define the mystery function above, write a one-line dictionary comprehension statement that achieves the same effect when applied to a dictionary `d` that has string keys and integer values.  

<!-- grade_id: cell-df433bce48997a2d, points: 2, grade: true, solution: true -->
```python
# using the same d defined in part (b) 

{k.upper() if v//2 == v/2 else k.lower():v for (k,v) in d.items() if v !=0}
```

Output:
```
{'AAA': 4, 'bb': 9, 'CCC': 20}
```